<h3>CITS2402 Lab 2</h3>

# Covid Tests per Capita: Is the US Leading the World?

<div>
    <img src="trump-testing.jpg" width=800/>
</div>

<br>

Former President Trump repeatedly stated that the US tests more than any other country in the world by far, and sometimes more than all the other countries put together.

In this sequel to the *Fake News!* case study, we'll investigate the veracity of one aspect of these claims, the *per capita* test rates.

In this lab we will also practise with:
- The Data (Science) Lifecycle - data acquisition, cleaning and conversion
- Working with (parsing) csv files
- Working with text and numbers (casting)
- Working with dictionaries

Before we get started, some details on the Study Buddy system.

---

## The "Study Buddy" System

As discussed in the lecture, the Study Buddy system is being trialled this year to allow students to work in pairs or triples (without triggering a plagiarism match) where desired. This is intended to strike a balance between:

* the benefits of peer learning,
* the need for students to be able to program individually (not least of which for subsequent tests).

To participate in the study buddy system, you just need to declare who you worked with on that lab, by putting the following in a *code cell*  near the top of your lab:

```
StudyBuddies = ['12345678','87654321']
```
The list of student numbers should include your own (optional) and up to two others. These should correspond to a similar declaration on (each of) the others' labs, in order to indicate that all have agreed to this grouping.

---


## Part 1

### Data Acquisition

#### *Our World in Data*

In the videos for this case study we obtained data, originally from the European CDC, through *Our World in Data*: https://ourworldindata.org/.

* Read the *[About](https://ourworldindata.org/about)* page to find out what this initiative is, what they hope to achieve, and who it is backed by. It is always important to know who is behind a data source in order to make an assessment about what degree of credibility, and potentially bias, to attribute to the source.

- Follow the link to `Health | Coronavirus Pandemic` and then to `Testing`. 

Note that the Data Scientists behind this site have provided a number of different ways of looking at and interpreting test rates. There is also a lot of background provided. 

* Scroll down to the section "Our checklist for COVID-19 testing data" and read through the ten items on the checklist.

This site is terrific example of Data Science done well!


### Per-capita testing

Find the section entitled "How many tests are performed each day".

Have a look at how the Map is presented, with the time slider that allows you to see snapshots over time.

Then look at the Chart view and scroll the cursor over the chart.

Again, these are great examples of interactive data presentation.

Note: Consistent with many of the media sites, we will refer to these data in general terms as the "per capita" data to distinguish them from totals or per case data. More precisely, however, they are tests per 1000 people. That is, they are 1000 times the per capita rate. The only reason for multiplying by 1000 is that its easier to read, say, 0.8 than 0.0008.

### Downloading and inspecting the data

- Go to `DOWNLOAD` and download the csv file `daily-tests-per-thousand-people-smoothed-7-day.csv`.

- Open the file in your favourite spreadsheet software and take a look. You should see a list of daily results.

- Have a look at the file in a terminal or your favourite text editor.

You should see, as anticipated from the filename extension, that this is a *comma separated values (csv)* file: in each row, the *fields* are separated by commas. There is also a header line, indicating what the data in each field represents.

Although various packages exist for reading in csv files (we'll visit some later), being proficient parsing data (breaking it into its constituent parts) is an important skill. In this lab we will therefore parse the file directly. The assessed questions therefore assume *no additional packages are imported*.

---


The current daily test file is currently over 4.6MB, and includes a lot of data since Trump's statements. We'll use a smaller version (430 KB), extracted on the 29th July 2020.

This has been distributed with the lab sheet, and can be found in the file `daily-tests-per-thousand-people-smoothed-7-day-20200729.csv`.

- Open this file in CoCalc and scan through it.

You should see that it is again in csv format with daily readings, but the way the data is reported has evolved over time. How does it compare with the latest download in terms of standardisation and consistency?

### Reading in the data

As usual, we'll start by setting up a constant with the path (in this case, its just the name) of the data file, so we don't need to keep typing it. We'll use the version of this file from the 29th July, distributed with the lab, so that we are all using the same file.

You can access this file with:

`DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"`

Note that unlike the previous labs sheets, in this lab we won't continue to put empty cells in for you to complete your code in. You can create cells as you need them using the '+' button (or shortcuts 'a' and 'b').



```python
DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"

```

```python
data = open(DATA, "r")
data1 = data.readlines()
for i in range(5):
    print(data1[i])
```

Output:
```
Entity,Code,Date,Daily tests per thousand people (7-day smoothed) (tests per thousand)

"Argentina, tests performed",,"Feb 18, 2020",0

Argentina,ARG,"Feb 18, 2020",0

Argentina,ARG,"Feb 19, 2020",0

Argentina,ARG,"Feb 20, 2020",0
```

* Read and print out the first 5 lines of data. The output should start like this:

```
Entity,Code,Date,Daily tests per thousand people (7-day smoothed) (tests per thousand)

Argentina,ARG,"Feb 18, 2020",0

Argentina,ARG,"Feb 19, 2020",0

Argentina,ARG,"Feb 20, 2020",0

Argentina,ARG,"Feb 21, 2020",0
```

In the first part of this lab we'll proceed on the assumption that the format at the start of the file is consistent throughout the file. (For this file you can see by eyeballing the data that it isn't. In very large data files, however, you cannot always tell by eyeballing, and you have to test your assumptions using code. We'll do this in Part 2.)



### Data Cleaning and Conversion

#### From strings to lists

* Read the first 5 lines again. This time use the `split` method to turn each line into a list before printing it out.

Your output should start like this:

```
['Entity', 'Code', 'Date', 'Daily tests per thousand people (7-day smoothed) (tests per thousand)\n']
['Argentina', 'ARG', '"Feb 18', ' 2020"', '0\n']
['Argentina', 'ARG', '"Feb 19', ' 2020"', '0\n']
```



```python
data = open(DATA, "r")
data1 = data.readlines()
list1=[]
for i in data1:
    line = i.split(",")
    line2 =[]
    for j in line:
        # j1 = j.strip()
        # line2.append(j1)
        line2.append(j)
    list1.append(line2)
for i in range(3):
    print(list1[i])
```

Output:
```
['Entity', 'Code', 'Date', 'Daily tests per thousand people (7-day smoothed) (tests per thousand)\n']
['"Argentina', ' tests performed"', '', '"Feb 18', ' 2020"', '0\n']
['Argentina', 'ARG', '"Feb 18', ' 2020"', '0\n']
```

Notice that we still have the newline character in the last item.

* Use the `strip` function to remove whitespace before splitting the lines.

Your output should now start like this:

```
['Entity', 'Code', 'Date', 'Daily tests per thousand people (7-day smoothed) (tests per thousand)']
['Argentina', 'ARG', '"Feb 18', ' 2020"', '0']
['Argentina', 'ARG', '"Feb 19', ' 2020"', '0']
```



```python
data = open(DATA, "r")
data1 = data.readlines()
list1=[]
for i in data1:
    i1 = i.strip()
    line = i1.split(",")
    line2 =[]
    for j in line:
        line2.append(j)
    list1.append(line2)
for i in range(3):
    print(list1[i])
```

Output:
```
['Entity', 'Code', 'Date', 'Daily tests per thousand people (7-day smoothed) (tests per thousand)']
['"Argentina', ' tests performed"', '', '"Feb 18', ' 2020"', '0']
['Argentina', 'ARG', '"Feb 18', ' 2020"', '0']
```

You may have noticed another problem, caused by the fact that the dates include commas. In fact, you may conclude that with this date format, the choice of `csv` and a comma as the *delimiter* (separator) was not a particularly good one, and an alternative such as tab separated (tsv) would have been better choice for these data.

Nevertheless, it is not ambiguous (importantly), because commas that are not intended as delimiters only appear within double quotes. The 'user' or *parser* of the file (our code in this case - but it could be, say, your favourite spreadsheet software) is expected to take the quotes into account when splitting the lines. 

There are many ways to deal with dates, but for now, we can just remove the comma and the quotes, since neither provide us with any information. The fields within the quotes (month, day and year) can be distinguished by their order (the comma is just for human consumption) and the quotes are otherwise redundant since it is a text file and all the fields will be read as strings.

Change your code so that it has a *preprocessing* step before it splits the lines. For each line after the header line your preprocessing step should:

  * find the positions of the two double quotes
  * extract the text and remove the quotes and the comma
  * replace the old date with the new date without quotes or a comma
  * throw a `ValueError` if the line doesn't have double quotes

We'll break it down into steps.



* Print each line (of the first 5, other than the header, and with the trailing whitespace removed) followed by the indices of the two double quotes:

```
Argentina,ARG,"Feb 18, 2020",0
14 27
Argentina,ARG,"Feb 19, 2020",0
14 27
Argentina,ARG,"Feb 20, 2020",0
14 27
...
```

Tip: You can put the double quote character in a string by enclosing it in single quotes (`'"'`).

Hint: Find the description of the String methods in the Python Library documentation, and compare the `find` and `index` methods. Why would you choose one or the other for this application?



```python
data = open(DATA, "r")
data1 = data.readlines()
list1=[]
for i in data1:
    i1 = i.strip()
    line = i1.split(",")
    line2 =[]
    for j in line:
        line2.append(j)
    list1.append(line2)

list1 = list1[1:]

for i in range(5):
    liststr = ",".join(list1[i])
    print(liststr)
    res = [j for j in range(len(liststr)) if liststr.startswith('"', j)]
    for j in res:
        print(j,"",end='')
    print()
```

Output:
```
"Argentina, tests performed",,"Feb 18, 2020",0
0 27 30 43 
Argentina,ARG,"Feb 18, 2020",0
14 27 
Argentina,ARG,"Feb 19, 2020",0
14 27 
Argentina,ARG,"Feb 20, 2020",0
14 27 
Argentina,ARG,"Feb 21, 2020",0
14 27
```

* Next, using the indices, print the line followed by the date string:

```
Argentina,ARG,"Feb 18, 2020",0
"Feb 18, 2020"
Argentina,ARG,"Feb 19, 2020",0
"Feb 19, 2020"
Argentina,ARG,"Feb 20, 2020",0
"Feb 20, 2020"
```



```python
data = open(DATA, "r")
data1 = data.readlines()
list1=[]
for i in data1:
    i1 = i.strip()
    line = i1.split(",")
    line2 =[]
    for j in line:
        line2.append(j)
    list1.append(line2)

list1 = list1[1:]

for i in range(5):
    liststr = ",".join(list1[i])
    print(liststr)
    res = [j for j in range(len(liststr)) if liststr.startswith('"', j)]
    liststr1 = liststr[res[0]:res[1]+1]
    print(liststr1)


```

* Now do the same, except with the quotes and any commas removed from the date string:

```
Argentina,ARG,"Feb 18, 2020",0
Feb 18 2020
Argentina,ARG,"Feb 19, 2020",0
Feb 19 2020
Argentina,ARG,"Feb 20, 2020",0
Feb 20 2020
```



```python
input_file = open(DATA, "r")
lines = input_file.readlines()
line_list=[]
for line in lines:
    l = line.strip()
    spe = l.split(",")
    line2 =[]
    temp_list = []
    found = False
    for j in line:
        if(j.startswith('"')):
            temp_list.append(j)
            found = True
        else:
            if(found):
                temp_list.append(j)
                if(j.endswith('"')):
                    line2.append(",".join(temp_list))
                    found = False
                    temp_list = []
            else:
                line2.append(j)
        line_list.append(line2)

list1 = list1[1:]

for i in range(5):
    liststr = ",".join(list1[i])
    print(liststr)
    res = [j for j in range(len(liststr)) if liststr.startswith('"', j)]
    liststr1 = liststr[res[0]+1:res[1]]
    liststr1 = liststr1.replace(",","")
    print(liststr1)


```

* Next, print the original lines with the old date field replaced by the cleaned date field:

```
Argentina,ARG,"Feb 18, 2020",0
Argentina,ARG,Feb 18 2020,0
Argentina,ARG,"Feb 19, 2020",0
Argentina,ARG,Feb 19 2020,0
Argentina,ARG,"Feb 20, 2020",0
Argentina,ARG,Feb 20 2020,0
```



```python
data = open(DATA, "r")
data1 = data.readlines()
list1=[]
for i in data1:
    i1 = i.strip()
    line = i1.split(",")
    line2 =[]
    for j in line:
        line2.append(j)
    list1.append(line2)

list1 = list1[1:]

for i in range(155, 164):
    liststr = ",".join(list1[i])
    print(liststr)
    res = [j for j in range(len(liststr)) if liststr.startswith('"', j)]
    liststr1 = liststr[res[0]+1:res[1]]
    liststr1 = liststr1.replace(",","")
    list2 = list1[i]
    list2[2] = liststr1
    print(",".join(list2))

```

#### 1. Checked Solution [2 lab marks]. 

Now that we have this working, we don't want this preprocessing step 'muddying' up our code, so let's put it together in a separate function. We'll also clean up the header line. We'll do this by writing a function according to the following *specification*.

Note that this specification is quite - well - specific! Please read it carefully. Irrespective of whether these are the choices you would have made at this point, it has to be precise because we need to know what will be returned under what circumstances, and when it will throw an error. 

* Write a function `clean (data_row)` that takes as its argument a string, and:
  * strips whitespace characters from the ends
  * if the string contains a double quote (`"`), strips the _first pair_ of double quotes and any commas between them
    * throws a ValueError if the first double quote is not paired with a second (this indicates an error in the data - you do not need to worry about what the error message is at this stage)
  * if the string doesn't contain any double quotes (eg. the header line), it truncates the line from the first opening parenthesis, `(`, if one exists, onwards
    * does not throw a ValueError if the line does not contain an opening parenthesis (we don't really care if an explanation in parentheses was in the header)

If called with the following code:

```
with open(DATAFILE,'r') as file:
    for i in range(5):
        print(clean(file.readline()))
```

the output should start like this:

```
Entity,Code,Date,Daily tests per thousand people
Argentina,ARG,Feb 18 2020,0
Argentina,ARG,Feb 19 2020,0
...
```

(Remember to remove `raise NotImplementedError()` and replace it with your code.)



<!-- grade_id: cleanq, points: 0, grade: false, solution: true -->
```python
def clean (data_row):
    data_rowlst=[]
    while data_row:
        data_rowlst.append(data_row)
        data_row = data.readline()

    list1=[]
    for i in data_rowlst:
        line = i.split("/n")
        line2 =[]
        for j in line:
            j1 = j.strip()
            line2.append(j1)
        list1.append(line2)
    listhead1 = "".join(list1[0])
    final = []
    final.append(listhead1[:49])
    
    for i in range(1,len(list1),1):
        liststr = "".join(list1[i])
        res = [j for j in range(len(liststr)) if liststr.startswith('"', j)]
        date = liststr[int(res[0])+1:int(res[1])].replace(",","")
        final.append(liststr[:res[0]-1] +","+ date + liststr[res[1]+1:])
    final2 = "/n".join(final)
    return final2


```

<!-- grade_id: clean_1, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal, assert_raises, raises
DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"
with open(DATA,'r') as file:
    assert_equal(clean(file.readline()), "Entity,Code,Date,Daily tests per thousand people ")
    assert_equal(clean(file.readline()), "Argentina,ARG,Feb 18 2020,0")
assert_raises(ValueError, clean, 'text"moretext')

print("So far, so good on the practice tests. Remember there will be additional tests applied. You should test your code to ensure it meets the specification.")

```

<!-- grade_id: clean_2, points: 1, grade: true, solution: false -->
```python
# This cell for marking use only.

```

#### From a file to a list

Next, rather than printing the lines, we'll store them all in a list (of lists).

* Define a variable `data_lists` as an empty list (). Write code that cleans and splits the (entire) input into lists (using the comma as a delimiter), and appends them to `data_lists`.

For example, the following should print the first five rows as a list of lists:

```
print(data_lists[:5])

[['Entity', 'Code', 'Date', 'Daily tests per thousand people'], ['Argentina', 'ARG', 'Feb 18 2020', '0'], ['Argentina', 'ARG', 'Feb 19 2020', '0'], ['Argentina', 'ARG', 'Feb 20 2020', '0'], ['Argentina', 'ARG', 'Feb 21 2020', '0']]
```

How many entries (lines of data) are there in the file?



#### Checking our cleaning so far

We now have a tidy list of lists, each with the four fields. Or do we?

With big data it may not be possible to manually look at every entry to see if we've accounted for every possibility. We should try to make our cleaning or preprocessing as general as possible so that we catch unexpected variations or bad data. 

In practice, we often have to make some assumptions about the data. However we should endeavour to test these.

In this case, we've assumed that the patterns we see at the start of the file continue through the file. So let's check that assumption.

* Write code that checks whether all your entries have 4 fields.

If not, why not? What have we missed?

Hint: Print out the first row (if there is one) where this is not true. Print out the number of that row, open the data file in CoCalc, and have a look at the data in that row. What do you find?


## Part 2

#### General vs Specific

The variations you see in the data are not unusual - remember that this is the collated official government data, its not an 'exercise'. It is *exactly the kind of thing you deal with as a working Data Scientist*. The structure of data is rarely decided upon and stuck to over time. This is not just due to lack of foresight - the use of the data, and therefore what is stored, often changes over time.

We'll need to come back to some differences in the content of the data, but for now let's focus on the formatting. Or, more precisely, *transforming* the data from the format in which it is provided to a format that is suitable for our use.

* Write a new function `cleaner()` to do more cleaning as necessary so that it is transformed into a list of lists, each with four fields.

You should try to make your code as **general as possible**. This means that, rather than just adjust for the specific case, think about *patterns*. 

For example, we have seen that the data format uses double quotes around fields containing the delimiter (a comma in this case). If the data is not corrupt (which is an assumption) therefore, we would expect the double quotes to always occur in pairs.
By focussing only on dates, we have been more *specific* than we need to be, so we may miss other cases. A more *general* solution will assume that the same pattern may occur in other fields.

Ensure you re-test your code after any changes you make to ensure it satisfies the requirements. (You might find it useful to write them all down.)

#### Type casting

Finally, the last field, tests per 1000 people, should be a `float`.

* Alter your code to change the tests ratio to a float.

Your first few lines should now look like this:
```
[['Entity', 'Code', 'Date', 'Daily tests per thousand people'], ['Argentina', 'ARG', 'Feb 18 2020', 0.0], ['Argentina', 'ARG', 'Feb 19 2020', 0.0], ['Argentina', 'ARG', 'Feb 20 2020', 0.0], ['Argentina', 'ARG', 'Feb 21 2020', 0.0]]
```

Again, we could make an assumption that the fourth string is always able to be converted to a float, but its possible that somewhere in the file that is not true. Later we will deal with this using *Exceptions*. For now, it is good practice to do some checks before attempting to cast. 
* Before casting, check that the fourth field is a non-empty string containing only decimal characters (as defined by Unicode) and up to one decimal point.

Tip: You may find the string function `isdecimal()` useful.

Remember to run your own tests, not rely on my tests or my sample output. For example, I've included the first few lines of output (the file is too large to include them all) but as we've seen those lines may not be indicative of the file as a whole. For one thing, they are all cases where the daily test ratio is zero.

Also, as always remember to check your functions with a "clean" kernel.

Remember that a Data Scientist is like a detective - always thinking about what we could possibly have missed, and testing for it.

Now let's put it all together in a more general function...

#### 2. Checked Solution [3 marks]

* Complete the function `get_cleaned_lists (raw_data)` so that it takes raw data, as read from the data file, and returns a list of lists, each containing the fields from the data file, with 
  * any pairs of quotes, and the commas within them, removed
    * mismatched quotes should throw a ValueError
  * any pairs of parentheses and text contained within them removed
    * mismatched parentheses should throw a ValueError (*note this is different to the earlier specification*)
  * any leading/trailing whitespace characters removed from all resulting fields
  * the daily tests ratio converted to a float, providing it consists of decimal characters and up to one decimal point

For example:
```
with open(DATAFILE,'r') as file:
    raw_data = file.read()
clean_data = get_cleaned_lists(raw_data)
print(clean_data[0])


['Entity', 'Code', 'Date', 'Daily tests per thousand people']
    

<!-- grade_id: d2098c, points: 0, grade: false, solution: true -->
```python
def get_cleaned_lists (raw_data):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: get_cleaned_lists_1, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
DATA = "daily-tests-per-thousand-people-smoothed-7-day-20200729.csv"
with open(DATA,'r') as file:
    data = get_cleaned_lists(file.read())
assert_equal(data[0],['Entity', 'Code', 'Date', 'Daily tests per thousand people'])
assert_equal(data[1],['Argentina', 'ARG', 'Feb 18 2020', 0.0])
assert_equal(data[40],['Argentina', 'ARG', 'Mar 28 2020', 0.008])
assert_equal(data[180],['Argentina tests performed', '', 'Mar 10 2020', 0.0])
print("So far, so good. Remember there will be additional tests applied.")

```

<!-- grade_id: get_cleaned_lists_2, points: 1, grade: true, solution: false -->
```python
# For marking use only.

```

<!-- grade_id: get_cleaned_lists_3, points: 1, grade: true, solution: false -->
```python
# For marking use only

```

## Part 3

### Data Presentation  - Using Dictionaries

*For this part it is left to you to break down the task into subtasks and test them as you go.*

We want to be able to access the daily test ratio for a country on a given date without using loops.

Dictionaries provide *much* faster access to data by *hashing* the dictionary keys.

* Store the daily tests data in a dictionary of dictionaries. The outer dictionary should use the countries as keys. The inner dictionary should use the dates as keys.

For example, if my outer dictionary is called `country_dict`, then evaluating `country_dict["Australia"]` should return:
    
```
{'Mar 29 2020': 0.382,
 'Mar 30 2020': 0.396,
 'Mar 31 2020': 0.41,
 'Apr 1 2020': 0.425,
 'Apr 2 2020': 0.439,
 'Apr 3 2020': 0.453,
 'Apr 4 2020': 0.467,
 ...
```
and `country_dict["Australia"]["Jul 1 2020"]` should return:
```
1.824
```

How many "countries" are there?


* Print out each country, followed by its number of tests per 1000 people, on 1st July, 2020. [If the country doesn't have a reading for that day, it can simply be skipped.]

It will be clear from this that some countries have more than one entry - for example, "Poland" and "Poland people tested" will show as two separate "countries". While we could easily clean these up, we haven't looked closely enough at the sources of the data to determine the reason for the different figures, and whether one is better than the other, to have grounds for choosing one over the other (you may wish to follow this up). Therefore we need to leave them all in for now.

### Presenting rankings in a table

* Finally, write a function `print_ranking(date)` that takes a date (as a string) and prints a table of testing results for that date, ranked from highest testing rate to the lowest.

So, for example, `print_ranking("Jul 1 2020")` will start as follows:
    
```
Testing results for Jul 1 2020
8.032 	 Luxembourg
5.653 	 United Arab Emirates
5.055 	 Bahrain
...
```

### Checking the news

The Whitehouse statement about US testing at the top of this sheet was made on April 28th.

As a Data Scientist, you have the power to process the raw data and fact-check statements like this yourself!

What do you make of President Trump's statements *from a per capita perspective*? Was he correct?

How does that compare with more recently?


In a May 11 _Rose Garden briefing_ President Trump stated:

>  _"We’re testing more people per capita than South Korea, the United Kingdom, France, Japan, Sweden, Finland, and many other countries — and, in some cases, combined."_

The BBC's May 15th article [Coronavirus: President Trump’s testing claims fact-checked](https://www.bbc.com/news/world-us-canada-52493073) "fact-checks" this claim (Claim One).

* Modify your function to the signature `print_ranking(date, countries=[])` so that:
  * if `countries` is omitted, it still prints the table for *all* countries reporting on that day
  * if a list of countries is passed to the function, then it only prints the table for those countries

Print the table for the US and those six countries on 11th May.

Is the BBC's fact check for the 11th May borne out by these data?

Can you think of a possible reason for these discrepancies? (Hint: Should "we're testing" be interpreted as a rate or as a cumulative total?)

On 22nd June Newsweek, in [Why Trump Is Both Right and Wrong About U.S. Coronavirus Testing Numbers](https://www.newsweek.com/why-trump-right-wrong-about-us-coronavirus-testing-1512472), compares the US with Russia, Spain, Germany and Portugal on cumulative per capita figures. 

How does this compare with the picture you get for the daily rate at this date?

Congratulations - you can now get a job as a fact-checking journalist! 🕵🏼

&copy; Cara MacNish, Univeristy of Western Australia

```python
# string with commas
s = "Yes, they backed Steve, Emma, and Rohan in the meeting, Emma, and Rohan in the meeting."
# split string by comma
ls = s.split(",", maxsplit=1)
print(ls)
```